# Activity: Creating Visualizations

## Introduction

In this activity you will practice using Pandas functionality to create 
visualizations.

This activity will cover the following topics:
- Creating individual and multiple box plots in one visualization.
- Creating individual and multiple histograms in one visualization.
- Creating scatter plots.


```python
import pandas as pd

# Data from https://github.com/mwaskom/seaborn-data/blob/2b29313169bf8dfa77d8dc930f7bd3eba559a906/mpg.csv
df = pd.read_csv('mpg.csv')
```

#### Question 1

Create a box plot from `df` using Pandas. Assign the number of box plots that
are created in the visualization to the variable `n_box_plots`.

<!-- grade_id: cell-42805a0ffa8c278e, points: 0, grade: false, solution: true -->
```python
# Your code here

# your code here

```

<!-- grade_id: cell-24df988d26b9c638, points: 1, grade: true, solution: false -->
```python
# Question 1 Grading Checks

assert isinstance(n_box_plots, int), 'Did you assign a number (integer) to n_box_plots representing the number of box plots that were created?'
```

#### Question 2

Using Pandas, create a visualization with box plots for the columns `cylinders`
& `model_year`. What is the smallest and largest number label on the axis
(representing values)?

Assign a best approximation of the smallest number label that is visible on the axis to the variable 
`smallest_label` and the largest number label to the variable `largest_label`. 

<!-- grade_id: cell-639cad323476f857, points: 0, grade: false, solution: true -->
```python
# Your code here

# your code here
```

<!-- grade_id: cell-9ff73f2857ed882c, points: 1, grade: true, solution: false -->
```python
# Question 2 Grading Checks

assert isinstance(smallest_label, (int, float)), 'Did you assign smallest_label a number value?'
assert isinstance(largest_label, (int, float)), 'Did you assign largest_label a number value?'
```

#### Question 3

Using Pandas, create a visualization of a histogram for the `acceleration`
column.
What is the approximate number for the peak of this histogram?

Assign the approximate number for the peak to the variable `approx_peak_height`.

<!-- grade_id: cell-119809466c7aa88e, points: 0, grade: false, solution: true -->
```python
# Your code here

# your code here
```

<!-- grade_id: cell-2a415bb4b446814a, points: 1, grade: true, solution: false -->
```python
# Question 3 Grading Checks


assert isinstance(approx_peak_height, (int, float)), 'Did you assign approx_peak_height a number value?'
```

#### Question 4

Using Pandas, create a scatter plot of `acceleration` vs `mpg` (`acceleration` 
on the x-axis and `mpg` y-axis).
What is the largest number label on the y-axis (representing `mpg`)?

Assign a best approximate for the largest number label on the `mpg` axis that is visible on the axis
to the variable `largest_mpg_label`. 

<!-- grade_id: cell-e744f46b13335128, points: 0, grade: false, solution: true -->
```python
# Your code here

# your code here
```

<!-- grade_id: cell-898cdf14fc5fed10, points: 1, grade: true, solution: false -->
```python
# Question 4 Grading Checks

assert isinstance(largest_mpg_label, (int, float)), 'Did you assign largest_mpg_label a number value?'
```