# Activity: Using Pandas

## Introduction

In this activity you will practice using some of the basic functionality associated with Pandas.
This activity will cover the following topics:
- Creating a `DataFrame`
- Displaying `DataFrame` information
- Accessing column data
- Getting ranges of column data
- Creating Series from dictionaries
- Using the `iloc()` and `loc()` methods
- Getting data from multiple columns
- Getting rows


#### Question 1

Create a DataFrame called `df` from the given CSV file `student_data.csv` then using the `df` DataFrame, assign the `Name` column to a `Series` called `names`.


<!-- grade_id: cell-98375e57b88587bc, points: 0, grade: false, solution: true -->
```python
import pandas as pd
# Your code here

# your code here
```

<!-- grade_id: cell-d179546d994dce5f, points: 1, grade: true, solution: false -->
```python
# Question 1 Grading Checks

assert isinstance(df, pd.DataFrame), "Did you create a DataFrame called df?"
assert isinstance(names, pd.Series), "Did you assign the Name column in a variable called names?"

```

#### Question 2

Using the `df` DataFrame, assign the `Age` and `Math Grade` columns to a DataFrame called `age_math`.


<!-- grade_id: cell-0625d61ebce59851, points: 0, grade: false, solution: true -->
```python
# Your code here

# your code here
```

<!-- grade_id: cell-54d4a96060947870, points: 1, grade: true, solution: false -->
```python
# Question 2 Grading Checks

assert isinstance(age_math, pd.DataFrame), "Did you assign the Age and Math Grade columns to a variable called age_math?"

```

#### Question 3

Using the `.loc()` method, assign the `Age` and `Math Grade` columns for the first 30 rows of `df` to a variable called `first_thirty_loc`.


<!-- grade_id: cell-89956eaaf3a9c1a6, points: 0, grade: false, solution: true -->
```python
# Your code here

# your code here
```

<!-- grade_id: cell-618e7ac032208e13, points: 1, grade: true, solution: false -->
```python
# Question 3 Grading Checks

assert isinstance(first_thirty_loc, pd.DataFrame), "Did you correctly assign the first 30 rows to a variable called first_thirty_loc?"

```

#### Question 4

Get the even numbered rows from the `Name` and `English Grade` columns and assign the result to a variable called `even_rows_english`.


<!-- grade_id: cell-e972ac5587c71643, points: 0, grade: false, solution: true -->
```python
# Your code here

# your code here
```

<!-- grade_id: cell-c1342bc37cb13e08, points: 1, grade: true, solution: false -->
```python
# Question 4 Grading Checks

assert isinstance(even_rows_english, pd.DataFrame), "Did you correctly assign even numbered rows to a variable called even_rows_english? Hint: the first row is index 0."

```