# Activity: Aggregations

## Introduction

In this activity you will practice using Pandas functionality to work with various aggregations.

This activity will cover the following topics:
- Measure an aggregate statistic over a specific column.
- Measure an aggregate statistic over a specific column of subsets using `groupby()` over one column.
- Measure multiple aggregate statistics over a specific column of subsets using `groupby()` over one column.
- Measure multiple aggregate statistics over a specific column of subsets using `groupby()` over multiple columns.
- Measure multiple aggregate statistics over a specific column of subsets using `groupby()` over multiple columns.
- Take the transpose of a DataFrame.


```python
import pandas as pd

# Data from https://github.com/mwaskom/seaborn-data/blob/2b29313169bf8dfa77d8dc930f7bd3eba559a906/mpg.csv
df = pd.read_csv('mpg.csv')
```

#### Question 1

Assign the average of the `weight` column where the `origin` is `'usa'` to the variable `weight_usa`.

<!-- grade_id: cell-f1c164ccdcccc694, points: 0, grade: false, solution: true -->
```python
# Your code here

# your code here
```

<!-- grade_id: cell-7d42f27a55c3269d, points: 1, grade: true, solution: false -->
```python
# Question 1 Grading Checks

assert isinstance(weight_usa, float), 'Did you assign a number to `weight_usa`?'
```

#### Question 2

Assign the maximum of the `weight` column where the `origin` is `'japan'` to the variable `weight_japan`.

<!-- grade_id: cell-332f4a0fe1fe253f, points: 0, grade: false, solution: true -->
```python
# Your code here

# your code here
```

<!-- grade_id: cell-7fd5bd50cbc0982e, points: 1, grade: true, solution: false -->
```python
# Question 2 Grading Checks

print(weight_japan)
```

#### Question 3

Using Pandas' `groupby()` method, group by the `cylinders` column to find the minimum and maximum `horsepower` and assign the result to the variable `weight_by_cylinder`.

<!-- grade_id: cell-a35bbf8bc6994cc6, points: 0, grade: false, solution: true -->
```python
# Your code here

# your code here
```

<!-- grade_id: cell-0a15e2627a69c686, points: 1, grade: true, solution: false -->
```python
# Question 3 Grading Checks

assert isinstance(weight_by_cylinder, pd.DataFrame), 'Did you create a DataFrame with `groupby()`?'
```