# Activity: Modifying & Replacing Values

## Introduction

In this activity you will practice modifying and replacing values in a DataFrame using the various method that Pandas has to offer.
This activity will cover the following, not necessarily in this order:
- Checking for anomalous values
- Using `.isnumeric()`
- Using `min()` and `max()` methods
- Using `.loc[]` to replace values
- Using `isnull()` and `notnull()` methods


#### Question 1

Create a `DataFrame` called `df` from the given CSV file `employee_data.csv`, and then create a mask called `valid_names` that checks the `Name` column for any non-numeric values.


<!-- grade_id: cell-36ca21c2f927b27e, points: 0, grade: false, solution: true -->
```python
import pandas as pd

# Your code here

# your code here
```

<!-- grade_id: cell-53f5cfdc632deeaa, points: 1, grade: true, solution: false -->
```python
# Question 1 Grading Checks

assert isinstance(df, pd.DataFrame), 'Have you created a DataFrame named df?'
assert isinstance(valid_names, pd.Series), 'Have you created a Series named valid_names?'

```

#### Question 2

Using the original `DataFrame` `df`, create a mask called `unknown_position` that checks the `Position` column for any values that are equal to the string `Unknown`. Then, replace all such values with `Engineer`.


<!-- grade_id: cell-91b265a8b1032f10, points: 0, grade: false, solution: true -->
```python
# Your code here

# your code here
```

<!-- grade_id: cell-8a77d3a0c913b652, points: 1, grade: true, solution: false -->
```python
# Question 2 Grading Checks

assert isinstance(unknown_position, pd.Series), 'Have you created a Series named unknown_position?'

```

#### Question 3

Using the original `DataFrame` `df`, create a mask called `invalid_vacation` that checks the `Weeks of Vacation` column for any values that are null or missing. Then, use that mask to assign the value 0 to them.


<!-- grade_id: cell-130ad98966b0ae41, points: 0, grade: false, solution: true -->
```python
# Your code here

# your code here
```

<!-- grade_id: cell-1093efc01f6c679a, points: 1, grade: true, solution: false -->
```python
# Question 3 Grading Checks

assert isinstance(invalid_vacation, pd.Series), 'Have you created a Series named invalid_vacation?'


```