# Activity: Selective Subsets

## Introduction

In this activity you will practice selecting subsets of data from a DataFrame using Pandas.
This activity will cover the following topics:
- Creating masks
- Negating masks
- Masks with slicing
- Null value masks


#### Question 1

Create a `DataFrame` called `df` from the given CSV file `movie_data.csv`, and then create a mask called `before_millennium` to select all movies that were released before 2000.


<!-- grade_id: cell-240a9f6b83e2bc8e, points: 0, grade: false, solution: true -->
```python
import pandas as pd

df = pd.read_csv("movie_data.csv")
before_millennium = df['Year Released'] < 2000

```

<!-- grade_id: cell-c67cccfe6b6864c5, points: 1, grade: true, solution: false -->
```python
# Question 1 Grading Checks

assert isinstance(df, pd.DataFrame), 'Did you create a DataFrame called df?'

```

#### Question 2

Using the `before_millennium` mask from Question 1, assign the titles of every movie that was released after 2000 to a `Series` called `newer_titles`.


<!-- grade_id: cell-a008246edc6867e8, points: 0, grade: false, solution: true -->
```python
newer_titles = df.loc[~before_millennium,'Title']

```

<!-- grade_id: cell-60a9d4c511c889a5, points: 1, grade: true, solution: false -->
```python
# Question 2 Grading Checks

assert isinstance(newer_titles, pd.Series), 'Did you create a Series called newer_titles?'

```

#### Question 3

Create a mask to select movies with a `Rating` of `8.9` and a `Box Office ($M)` value higher than `1000.0`. Assign the resulting `Series` to a variable called `popular_pg_movies`.


<!-- grade_id: cell-61027e670582168c, points: 0, grade: false, solution: true -->
```python
popular_pg_movies = df.loc[(df['Rating'] == 8.9) & (df['Box Office ($M)'] > 1000.0), 'Title']
popular_pg_movies
```

Output:
```
19    The Lord of the Rings: The Return of the King
Name: Title, dtype: object
```

<!-- grade_id: cell-e29d54428ab17eed, points: 1, grade: true, solution: false -->
```python
# Question 3 Grading Checks

assert isinstance(popular_pg_movies, pd.Series), 'Did you create a Series called popular_pg_movies?'

```

#### Question 4

Create a mask to select movies with a null value for `Box Office ($M)` or `Rating`. Assign the resulting `Series` to a variable called `missing_info`.


<!-- grade_id: cell-ea4f4905c6892e32, points: 0, grade: false, solution: true -->
```python
# Your code here

# your code here


```

<!-- grade_id: cell-543bf8c0aea1a245, points: 1, grade: true, solution: false -->
```python
# Question 4 Grading Checks

assert isinstance(missing_info, pd.Series), 'Did you create a Series called missing_info?'

```